# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest dtring frmo the liset of input strings and restursn the position number of that longest string.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['one', 'two', 'three', 'four'])
    2
    """
    longest_pos = 0
    longest_len = 0
    
#    print("Position  Item       Longest So Far")
#    print("--------  ---------  -------------")
    for position, item in enumerate(strings):
        
#        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
            
    return longest_pos

### END SOLUTION
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

```python
def longest(L):
    """(list) -> int
    Finds the longest string in this list.
    
    >>> longest(['my','name','is','Paul Boal'])
    3
    
    >>> longest("happy birthday to me".split(" "))
    1
    """
    
    position = 0
    length = 0
    
    for pos, word in enumerate(L):
        if len(word) > length:
            position = pos
            length = len(word)
    
    return position
    
    
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['my','name','is','Paul Boal'])
Expecting:
    3
ok
Trying:
    longest("happy birthday to me".split(" "))
Expecting:
    1
ok
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION

def f_to_c(temp_f):
    """(float) -> float
    Converts Farenheit to Celsius
    
    >>> f_to_c(212)
    100.0
    
    >>> f_to_c(32)
    0.0
    
    >>> f_to_c(98.6)
    37.0
    """
    
    return (5/9 * (temp_f - 32))
    
### END SOLUTION
```

```python
### BEGIN SOLUTION
def f_to_c(F):
    """(float) -> float
    
    >>> f_to_c(212)
    100.0
        
    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    return (5/9) * (F-32)

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


### BEGIN SOLUTION
def los(admit, discharge, diagnosis):
    """(date, date, string) -> int
    Return length of stay in days based on the rules provided above.
    If diagnosis is "Observation" then LOS = 0.
    Otherwise LOS is number of whole days between admit and discharge dates.
    
    >>> los(date(2019,1,3), date(2019,1,8), 'COPD')
    5
    
    >>> los(date(2020,1,5), date(2020,1,8), 'Observation')
    0
    """
    
    d = (discharge - admit).days
    if diagnosis == 'Observation':
        return 0
    else:
        return d
    
### END SOLUTION
```

```python
def los(admit, discharge, diagnosis):
    """(date, str, date)->int
    Computes length of stay in days.  If diagnosis is "Observation" then set the length of stay to 0.

    >>> los(date(2021,1,3), date(2021,1,8), 'COPD')
    5
    
    >>> los(date(2021,1,5), date(2021,1,9), 'Hypertension')
    4
    
    >>> los(date(2021,1,23), date(2021,1,25), 'Observation')
    0
    """
    if diagnosis == "Observation":
        return 0
    else:
        return (discharge - admit).days
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2021,1,3), date(2021,1,8), 'COPD')
Expecting:
    5
ok
Trying:
    los(date(2021,1,5), date(2021,1,9), 'Hypertension')
Expecting:
    4
ok
Trying:
    los(date(2021,1,23), date(2021,1,25), 'Observation')
Expecting:
    0
ok
```

```python
for e, v in encounters.items():
    print(los(v['admit'],v['discharge'],v['diagnosis']))
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1233": { "admit": date(2019,1,4), "diagnosis": "COPD", "discharge": date(2019,1,9) },
        "E1232": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,8) },
        "E1231": { "admit": date(2019,1,2), "diagnosis": "Observation", "discharge": date(2019,1,8) }
}

### BEGIN SOLUTION
def los2(admit, discharge, diagnosis):
    """(date, date, string) -> int
    Return length of stay in days based on the rules provided above.
    If diagnosis is "Observation" then LOS = 0.
    Otherwise LOS is number of whole days between admit and discharge dates.
    
    >>> los(date(2019,1,3), date(2019,1,8), 'COPD')
    5
    
    >>> los(date(2020,1,5), date(2020,1,8), 'Observation')
    0
    """
    
    d = (discharge - admit).days
    return d

def average_los(encounters):
    """(dict) -> float
    This function will return the average length of stay from data provided in the dictionaries.
    If the LOS for an individual encounter is 0, then it is not counted to the average.
    
    >>> average_los(test1)
    5
    
    >>> average_los(test2)
    4.5
    """
    
    total_los = 0
    total_encounters = len(encounters)
    
    for enc in encounters:
        enc_los = los2(encounters[enc].get('admit'), encounters[enc].get('discharge'), encounters[enc].get('diagnosis'))
        if enc_los != 0:
            total_los += enc_los
        else:
            total_encounters -= 1
    
    if total_encounters == 0:
        return 0
    
    return total_los / total_encounters

### END SOLUTION
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    average_los(test1)
Expecting:
    5
**********************************************************************
File "__main__", line 33, in NoName
Failed example:
    average_los(test1)
Expected:
    5
Got:
    5.0
Trying:
    average_los(test2)
Expecting:
    4.5
**********************************************************************
File "__main__", line 36, in NoName
Failed example:
    average_los(test2)
Expected:
    4.5
Got:
    5.0
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION

def c_to_f(temp_c):
    """(float) -> float
    Converts from Celsius to Farenheit
    
    >>> c_to_f(100)
    212.0
    
    >>> c_to_f(0)
    32.0
    
    >>> c_to_f(37)
    98.60000000000001
    """

    return (temp_c * (9/5)) + 32
    
### END SOLUTION
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    c_to_f(37)
Expecting:
    98.60000000000001
ok
```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-19-cef8011cb395> in <module>
----> 1 assert False, "DO NOT REMOVE THIS LINE"

AssertionError: DO NOT REMOVE THIS LINE
```

---

```python
%%bash
git pull
git add week05_assignment_2.ipynb
git commit -a -m "Submitting the week 5 programming assignment"
git push
```

Output:
```
Already up to date.
[main 5c4b9a7] Submitting the week 5 programming assignment
 2 files changed, 755 insertions(+), 2 deletions(-)
 create mode 100644 week05/week05_assignment_2.ipynb
To github.com:shreyadash/hds5210-2021.git
   bde67c6..5c4b9a7  main -> main
```


---

If the message above says something like _Submitting the week 5 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```