# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 15 through 20 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names separated by commas.

```
sort_names([
['Paul', 'Boal'],
['Kermit','the','Frog'],
['Donald','Duck']])

'Paul Boal, Donald Duck, Kermit the Frog'
```

```python
names = [['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck']]
```

```python
formatted_names = []

for name in names:
    if len(name) == 3:
        sortable_name = name[2] + ', ' + name[1] + ' ' + name[0]
        printable_name = name[0] + ' ' + name[1] + ' ' + name[2]
    else:
        sortable_name = name[1] + ', ' + name[0]
        printable_name = name[0] + ' ' + name[1]
    formatted_names.append([sortable_name, printable_name])
    
formatted_names.sort()
print(formatted_names)

output = ''
for name in formatted_names:
    if output == '':
        output = name[1]
    else:
        output += ', ' + name[1]

print(output)
```

Output:
```
[['Boal, Paul', 'Paul Boal'], ['Duck, Donald', 'Donald Duck'], ['Frog, the Kermit', 'Kermit the Frog']]
Paul Boal, Donald Duck, Kermit the Frog
```

```python
def sort_names(names):
    
    sorted_names = ""
    ### BEGIN SOLUTION
    # 1. Restructure list to have LAST FIRST MIDDLE
    # 2. I need to keep FIRST MIDDLE LAST along while I'm sorting
    #.     ['LAST FIRST MIDDLE', 'FIRST MIDDLE LAST']
    
    formatted_names = []

    for name in names:
        if len(name) == 3:
            sortable_name = name[2] + ', ' + name[1] + ' ' + name[0]
            printable_name = name[0] + ' ' + name[1] + ' ' + name[2]
        else:
            sortable_name = name[1] + ', ' + name[0]
            printable_name = name[0] + ' ' + name[1]
        formatted_names.append([sortable_name, printable_name])

    formatted_names.sort()

    for name in formatted_names:
        if sorted_names == '':
            sorted_names = name[1]
        else:
            sorted_names += ', ' + name[1]
    ### END SOLUTION
    
    return sorted_names
```

```python
sort_names([['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck']])
```

Output:
```
'Paul Boal, Donald Duck, Kermit the Frog'
```

```python
def sort_names(names):
    
    sorted_names = ""
    
    # 1. Create a list that I can easily sort
    # 2. Sort that list
    # 3. Format the sorted names the way I want them
    # 4. Combine name into an output string
    
    return sorted_names
```

```python
names = ['Paul', 'Boal'],['Kermit','the','Frog'],['Donald','Duck']

# 1. Create a list that I can easily sort
#    But keep a key to the original input so I can go back to that later
sortable_names = []
for key, name in enumerate(names):
    if len(name) == 3:
        sortable_name = '{}, {} {}'.format(name[2],name[0],name[1])
    else:
        sortable_name = '{}, {}'.format(name[1],name[0])
    sortable_names.append((sortable_name, key))

# 2. Sort that list
sorted_names = sorted(sortable_names)
sorted_names
```

Output:
```
[('Boal, Paul', 0), ('Duck, Donald', 2), ('Frog, Kermit the', 1)]
```

```python
# 3. Format that sorted names the way I want them
pretty_names = []
for name in sorted_names:
    name_key = name[1]
    name_parts = names[name_key]
    pretty_name = ' '.join(name_parts)
    pretty_names.append(pretty_name)

pretty_names
```

Output:
```
['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

```python
# 4. Combine names into an output string
output = ', '.join(pretty_names)
output
```

Output:
```
'Paul Boal, Donald Duck, Kermit the Frog'
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

```python
# Loop through the dictionary key and value
# For each item in dict, I know value is going to be a list of first names
#For each first name, I need to output [LAST, FIRST]

families = {
    'Boal': ['Paul', 'Anny', 'Jim'],
    'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
    'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

output = []
for last_name, first_names in families.items():
    for first in first_names:
        output.append([last_name, first])

print(output)
```

Output:
```
[['Boal', 'Paul'], ['Boal', 'Anny'], ['Boal', 'Jim'], ['Lester-Boal', 'Ellie'], ['Lester-Boal', 'Ada'], ['Lester-Boal', 'Teddy'], ['Lester', 'Sarahlynn'], ['Lester', 'Jessica'], ['Lester', 'Grace'], ['Lester', 'Carolynn']]
```

```python
def dict_to_list(families):
    
    output = []
    for last_name, first_names in families.items():
        for first in first_names:
            output.append([last_name, first])
            
    return output
```

```python
dict_to_list({'Duck': ['Hewey', 'Dewey', 'Louis'],
             'Mouse': ['Mickey', 'Minney']})
```

Output:
```
[['Duck', 'Hewey'],
 ['Duck', 'Dewey'],
 ['Duck', 'Louis'],
 ['Mouse', 'Mickey'],
 ['Mouse', 'Minney']]
```

```python
def dict_to_list(families):
    
    output = []
    for family, names in families.items():
        for name in names:
            output.append([family, name])
            
    return output
```

```python
dict_to_list({
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
})
```

Output:
```
[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Ellie'],
 ['Lester-Boal', 'Ada'],
 ['Lester-Boal', 'Teddy'],
 ['Lester', 'Sarahlynn'],
 ['Lester', 'Jessica'],
 ['Lester', 'Grace'],
 ['Lester', 'Carolynn']]
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict():

    names = [['Boal', 'Paul'],
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # To help you, I've provided a blank dictionary named `families`.
    families = {}
    for name in names:
        if name[0] in families:
            families[name[0]].append(name[1])
        else:
            families[name[0]] = [name[1]]
    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.

    ### BEGIN SOLUTION
    # Put your solution here
    ### END SOLUTION

    return families

```

```python
list_to_dict()
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict() == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):

    conclusions = []
    for patient in patients:
        los = avg_los.get(patient[1])
        if patient[2] == los:
            conclusions.append([patient[0],'just right'])
        elif patient[2] < los:
            conclusions.append([patient[0],'too short'])
        else:
            conclusions.append([patient[0],'too long'])
    ### BEGIN SOLUTION
    # Put your solution here
    ### END SOLUTION
    
    return conclusions
```

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

---

```python
%%bash
git pull
git add week04_assignment_2.ipynb
git commit -a -m "Submitting the week 4 programming assignment"
git push
```


---

If the message above says something like _Submitting the week 4 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.